// Copyright (C) 2020 Mike Pennisi authors. All rights reserved.
// This code is governed by the BSD license found in the LICENSE file.
/*---
esid: sec-comments
description: Conflict comment "end" with no optional comment characters
info: |
    MultiLineComment::
      /* MultiLineCommentCharsopt *\/
      <<<<<<< MultLineCommentCharsopt =======

    SingleLineComment::
      // SingleLineCommentCharsopt
      >>>>>>> SingleLineCommentCharsopt
---*/

>>>>>>> well-intentioned
